// variables1.rs
// Make me compile! Execute the command `rustlings hint variables1` if you want a hint :)

// About this `I AM NOT DONE` thing:
// We sometimes encourage you to keep trying things on a given exercise,
// even after you already figured it out. If you got everything working and
// feel ready for the next exercise, remove the `I AM NOT DONE` comment below.

fn main() {
    //变量需要使用let声明，赋值时已确定变量类型
    //也可使用mut在变量名前声明变量为可变值,即后续使用时可被修改
    let mut x = 5;
    println!("x has the value {}", x);
}
